### Week 5-2 NumPy Practice Problems

```python
import numpy as np
```

### Warm up
NumPy creation and indexing.

```python
# 1. Creates 3x3 array of zeros
np.zeros((3,3))
```

Output:
```
array([[0., 0., 0.],
       [0., 0., 0.],
       [0., 0., 0.]])
```

```python
# 2. Creates array [0, 2, 4, 6, 8]
np.arange(0,10,2)
```

Output:
```
array([0, 2, 4, 6, 8])
```

```python
# 3. 5 evenly spaced numbers from 0 to 1
np.linspace(0,1,5)
```

Output:
```
array([0.  , 0.25, 0.5 , 0.75, 1.  ])
```

```python
# Array indexing, Gets elements > 5
arr = np.array([[1, 2, 3], 
                [4, 5, 6], 
                [7, 8, 9]])

print(arr[:2, 1:2])

for i in arr: 
    for j in i: 
        if j >= 5: 
            print(j)

#BETTER!!

print(arr[arr > 5])
```

Output:
```
[[2]
 [5]]
5
6
7
8
9
[6 7 8 9]
```

### Problem 1
Convert an array of Celsius measurements to Fahrenheit.
Note: `fahrenheit = (Celsius * 9/5) + 32`.

```python
def convert_to_fahrenheit_list(temperatures):
    far = []
    for i in temperatures: 
        far.append((i * 9/5) + 32)
    return far
```

```python
print("Testing with a regular list:")
temps_list = [0, 10, 20, 30]
print(convert_to_fahrenheit_list(temps_list))
```

Output:
```
Testing with a regular list:
[32.0, 50.0, 68.0, 86.0]
```

```python
def convert_to_fahrenheit_numpy_array(temperatures):
    return (temperatures * 9/5) + 32
```

```python
print("\nTesting with a NumPy array:")
temps_array = np.array([0, 10, 20, 30])
print(convert_to_fahrenheit_numpy_array(temps_array))
```

Output:
```
Testing with a NumPy array:
[32. 50. 68. 86.]
```

### Problem 2
Use numpy to create a function that adds a border filled with zeros around an existing array.

*Hint*: Use numpy's `pad` function. See https://numpy.org/doc/stable/reference/generated/numpy.pad.html

Note: we did not cover this function in class, see if you are comfortable of reading numpy documents and learn on your own!

<!-- grade_id: cell-2466342b05a17106, points: 0, grade: false, solution: true -->
```python
import numpy as np
def add_border(x):
    ans = np.pad(x, ((1,1), (1,1)), 'constant', constant_values=0)
    return ans
```

<!-- grade_id: cell-fb9bae34e3908aed, points: 1, grade: true, solution: false -->
```python
x = np.ones((1, 1))
claimed = add_border(x)
correct = np.zeros((3, 3))
correct[1, 1] = 1.0
assert claimed.shape == correct.shape
assert np.all(np.abs(claimed - correct) < 1e-8)
```

### Problem 3
Create an $n \times n$ numpy array filled with a checkerboard pattern of zeros and ones. The first entry should be the integer `0`. Assume $n \ge 2$.

```python
def make_checkerboard(n):
    row = np.arange(n) #creates a row from 0 to n
    col = np.arange(n).reshape(n,1) #creates a row from 0 to n then reshapes it to nx1 column
    print(row)
    print(col)
    print(row + col)
    return (row + col) % 2 #add together, then divide by 2; if even, 0, if odd, 1.
```

```python
make_checkerboard(10)
```

Output:
```
[0 1 2 3 4 5 6 7 8 9]
[[0]
 [1]
 [2]
 [3]
 [4]
 [5]
 [6]
 [7]
 [8]
 [9]]
[[ 0  1  2  3  4  5  6  7  8  9]
 [ 1  2  3  4  5  6  7  8  9 10]
 [ 2  3  4  5  6  7  8  9 10 11]
 [ 3  4  5  6  7  8  9 10 11 12]
 [ 4  5  6  7  8  9 10 11 12 13]
 [ 5  6  7  8  9 10 11 12 13 14]
 [ 6  7  8  9 10 11 12 13 14 15]
 [ 7  8  9 10 11 12 13 14 15 16]
 [ 8  9 10 11 12 13 14 15 16 17]
 [ 9 10 11 12 13 14 15 16 17 18]]
array([[0, 1, 0, 1, 0, 1, 0, 1, 0, 1],
       [1, 0, 1, 0, 1, 0, 1, 0, 1, 0],
       [0, 1, 0, 1, 0, 1, 0, 1, 0, 1],
       [1, 0, 1, 0, 1, 0, 1, 0, 1, 0],
       [0, 1, 0, 1, 0, 1, 0, 1, 0, 1],
       [1, 0, 1, 0, 1, 0, 1, 0, 1, 0],
       [0, 1, 0, 1, 0, 1, 0, 1, 0, 1],
       [1, 0, 1, 0, 1, 0, 1, 0, 1, 0],
       [0, 1, 0, 1, 0, 1, 0, 1, 0, 1],
       [1, 0, 1, 0, 1, 0, 1, 0, 1, 0]])
```

<!-- grade_id: cell-e6f9eb1adf863ce2, points: 1, grade: true, solution: false -->
```python
correct = \
    [[0, 1, 0, 1, 0, 1, 0, 1],
     [1, 0, 1, 0, 1, 0, 1, 0],
     [0, 1, 0, 1, 0, 1, 0, 1],
     [1, 0, 1, 0, 1, 0, 1, 0],
     [0, 1, 0, 1, 0, 1, 0, 1],
     [1, 0, 1, 0, 1, 0, 1, 0],
     [0, 1, 0, 1, 0, 1, 0, 1],
     [1, 0, 1, 0, 1, 0, 1, 0]]
assert np.all(correct == make_checkerboard(8))
```

Output:
```
[0 1 2 3 4 5 6 7]
[[0]
 [1]
 [2]
 [3]
 [4]
 [5]
 [6]
 [7]]
[[ 0  1  2  3  4  5  6  7]
 [ 1  2  3  4  5  6  7  8]
 [ 2  3  4  5  6  7  8  9]
 [ 3  4  5  6  7  8  9 10]
 [ 4  5  6  7  8  9 10 11]
 [ 5  6  7  8  9 10 11 12]
 [ 6  7  8  9 10 11 12 13]
 [ 7  8  9 10 11 12 13 14]]
```

```python
import numpy as np
x = np.arange(10)
```