### Lecture 3-2 Practice: Lambda functions and Dictionaries

### Question 1
You have a list of 2D coordinates as tuples (x, y). Sort them by their distance from the origin (0, 0) using a lambda function.

<!-- grade_id: cell-6c67d579c481de9b, points: 0, grade: false, solution: true -->
```python
import math
coordinates = [(3, 4), (1, 1), (5, 0), (2, 3), (0, 2)]
sorted_coordinates = # YOUR CODE HERE                
print(sorted_coordinates)
# Expected output: [(1, 1), (0, 2), (2, 3), (3, 4), (5, 0)]
```

### Question 2
Define a function called `switch_key_value`, which takes a dictionary as its only argument. Switch the placement of key and value of the dictionary and return it as a new dictionary.

<!-- grade_id: cell-6c67d579c481de9b, points: 0, grade: false, solution: true -->
```python
def switch_key_value(d):
    raise NotImplementedError()           
```

<!-- grade_id: cell-a0b5ecba31ff8afb, points: 1, grade: true, solution: false -->
```python
assert(switch_key_value({}) == {})
assert(switch_key_value({'a': 1, 'b': 2}) == {1: 'a', 2: 'b'})
```

### Question 3
Define a function called `sort_list_dict`, which takes a dictionary whose values are lists as its only argument. Sort all these lists alphabetically return them in a new dictionary (with the same keys).

<!-- grade_id: cell-996c61e79b68ca6d, points: 0, grade: false, solution: true -->
```python
def sort_list_dict(d):
    raise NotImplementedError()           
```

<!-- grade_id: cell-f752b519fbf2ac78, points: 1, grade: true, solution: false -->
```python
assert(sort_list_dict({}) == {})

assert(sort_list_dict({'a': []}) == {'a': []})

input_dict = {'a': [4, 1, 2, 3], 'b': [3, 2, 1]}
desired_output = {'a': [1, 2, 3, 4], 'b': [1, 2, 3]}
assert(sort_list_dict(input_dict) == desired_output)

input_dict = {'S': [2, 3, 1], 'T': [5, 1, 2], 'D': [3, 2, 4]}
desired_output = {'S': [1, 2, 3], 'T': [1, 2, 5], 'D': [2, 3, 4]}
assert(sort_list_dict(input_dict) == desired_output)
```