# STATS 507, Week 1 practice

```python
def my_add(x,y):
    return x+y
print(my_add(3,4))
```

Output:
```
7
```

```python
import torch

map_lanes = torch.cat([
    torch.randn(2,11,48,100,5).float(),
    torch.randn(2,11,48,100,1).float(),
    torch.randint(0, 11, (2, 11, 48, 100, 1)).int(),
    torch.randint(0, 11, (2, 11,48,100, 1)).int(),
], dim=-1)
```

### Problem 1
Write a Python function that prints "Hello, World!".

```python
def fun():
    print("Hello, World!")

fun
```

Output:
```
Hello World!
```

### Problem 2
Write a function that takes two integers as argumens. Return the first argument if it is a positive number. Otherwise, return the sum of both arguments.

<!-- grade_id: cell-db7233a2d95c6406, points: 0, grade: false, solution: true -->
```python
def postive_or_sum(x, y):
    if x > 0: 
        return x
    else:
        return x+y

assert postive_or_sum(1,2) == 1
```

<!-- grade_id: cell-d525ecbed5402db1, points: 1, grade: true, solution: false -->
```python
assert postive_or_sum(42, 16) == 42
assert postive_or_sum(33, 13) == 33
assert postive_or_sum(-33, 13) == -20
```

```python
assert postive_or_sum(3, 4) == 3
```

```python
assertEqual
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[11], line 1
----> 1 assertEqual

NameError: name 'assertEqual' is not defined
```

```python

```